```python
# Initialize Otter
import otter
grader = otter.Notebook("week2.ipynb")
```

# HW2: Intermediate Programming and Algorithms
#### CS-31: Foundations for Artificial Intelligence
*Point breakdown: 40 autograded · 50 manually graded · 10 collaboration statement*

*You are allowed to work in groups for the lab. You must write anything that will be turned in - all code and all written solutions - on your own. You must list all people and resources that you interacted with in the Collaboration Statement below. See [course policies](https://www.cs.tufts.edu/comp/31/policy.html) for more information.*

---

In this assignment we'll be leveling up some of your programming skills in Python, as well as getting you to use data structures to reason more carefully about the efficiency and complexity of your code when writing programs. These topics are important to know both for understanding the underlying mechanics of the algorithms we will talk about in this class and how to reason about them - but also will make you a better programmer in general! These are things you'll see in production or research Python code that usually isn't taught in intro classes (both in AI, but also if you use Python anywhere else), so it's good to get familiar.

For due dates and submission instructions: please see the course website.

---

- Lists / Dict Comprehensions
- Reading JSON Files, Navigating Heavily Nested Data
- Advanced Functions 
- Classes and Objects
- Data Structures
    - Sets, Queues, and Stacks

Plus, you'll put that all together to improve the efficiency of a very simple decision algorithm.

### **Table of Contents**

1) List and Dictionary Comprehensions
2) I/O and Heavily Nested Data Structures
3) Other Python Potpourri
4) Classes and Objects
5) Data Structures and Complexity
6) Putting It All Together
7) Collaboration Statement

## **1. List and Dictionary Comprehensions**

This is more of a syntax trick to make things a little easier for you, but you'll see this a lot in professional Python code - so it's good to know!

[List comprehensions](https://www.w3schools.com/python/python_lists_comprehension.asp) offer a shorter syntax when you want to create a new list based on the values of an existing list. Basically, if you ever found yourself writing a bunch of short `for` statements that do very simple things (such as filtering data into a new list, or creating a copy of a list with a transformation applied), list comprehensions are a quicker way to do that.

For example, let's say I have a list of numbers, and I want to create a new list of numbers that multiplies each value in that list by two.

Here's one way you'd do it.

### **1.1 List Comprehensions**

```python
items = [2, 231, 424, 2525, 2312, 2313]

mult_2 = []
for i in items:
    new_item = i*2
    mult_2.append(new_item)
    
print(mult_2)
```

This type of pattern shows up so much that Python has a shorthand for it - it essentially squeezes the whole loop into one line.

```python
# The same code as above, but using a list comprehension.
mult_2_lc = [i*2 for i in items]
print(mult_2_lc)
```

Another example: say we have a list of names and only want to grab the names that start with the letter `A`.

```python
names = ["Bob", "Anne", "Azu", "Carl", "Zuzu", "Apple", "Sonia"]

starts_with_a = []
for name in names:
    if name[0] == "A":
        starts_with_a.append(name)

print(starts_with_a)
```

Here's the same code, but with a list comprehension. Notice how we add an `if` at the end.

```python
starts_with_a = [name for name in names if name[0] == "A"]
print(starts_with_a)
```

And of course, we can combine these!

```python
names = ["Bob", "Anne", "Azu", "Carl", "Zuzu", "Apple", "Sonia"]

starts_with_a_all_caps = []
for name in names:
    if name[0] == "A":
        new_name = name.upper()
        starts_with_a_all_caps.append(new_name)

print(starts_with_a_all_caps)
```

```python
starts_with_a_all_caps = [name.upper() for name in names if name[0] == "A"]
print(starts_with_a_all_caps)
```

#### **Exercise 1.1: Odd Squares.**

In a single line, use a list comprehension to create a list that contains the square of all the odd numbers in the list.

```python
nums = list(range(1,11))

odds_squared = ...

print(odds_squared)
```

```python
grader.check("q_odds_squared")
```

### **1.2 Dictionary Comprehensions**

Surprise, surprise - Python has a very similar thing for dictionaries!

Suppose I have an existing dictionary, and I want to create a copy of it where 1) each key is fully uppercased and 2) each value is squared.

```python
d = {
    "Bob": 1,
    "Alice": 2,
    "Zuzu": 4,
    "Blarg": 9
}

new_dict = {}
for key, value in d.items():
    new_dict[key.upper()] = value**2

print(new_dict)
```

And the same code in dictionary comprehension form...

```python
new_dict = {key.upper() : value**2 for key,value in d.items() }
print(new_dict)
```

**Exercise 1.2: Filtering a Dictionary.** 

Create a new dictionary that only has items with values less than 4.

You can look at [this reference](https://www.geeksforgeeks.org/python/python-dictionary-comprehension/) for a refresher. The syntax is very similar!

```python
d = {
    "Bob": 1,
    "Alice": 2,
    "Zuzu": 4,
    "Blarg": 9
}

new_dict = ...

print(new_dict)
```

```python
grader.check("q_dict_filter")
```

## **2. I/O and Heavily Nested Data Structures**

Everyone wants to do AI, but nobody wants to keep their data clean.

Aside from `.csv` (i.e. spreadsheets), the second type of data format you'll see a lot of in AI is `.JSON` files. 

Open the `airports.json` file that shipped with the assignment in a new tab in VSCode and see how it's structured. Unlike spreadsheets which have a fixed number of rows and columns, `.JSON` files allow arbitrary nesting of data. This is nice because it gives us a lot more flexibility to store different types of relationships that can't fit nicely into a spreadsheet - but also makes it a bit more difficult to work with.

Real world data is often lists of dictionaries with multiple fields, some of those fields being lists (which sometimes contain dictionaries) and... - you get the idea. So it's good to get proficiency writing code that can navigate these structures.

```python
import json

# Loading the file
with open("./airports.json") as f:
        airports = json.load(f)

# Let's take a peek
airports
```

```python
# Let's just look at the first entry and try to understand the structure.

# We can see that our airports have:
# - Location
# - City
# - Max Capacity
# - Current Occupancy

# We see that `forecast` itself is a list
# It seems that each `forecast` contains
#   - A temperature
#   - A rain level
#   - A list of "adverse_events", such as "Snow" or "High Winds"
airports[0]
```

#### **Exercise 2.1: Counting Adverse Events.**

Write a function that takes in an airport name, and returns a dictionary that counts the different types of adverse events present in that airport's forecast.

For example, `get_adverse_event_count(airports, "KBOS")` should return a dictionary like:

`{'Snow': 1, 'High Winds': 2, 'Wind Shear': 1}`

As there was one instance of `Snow`, two instances of `High Winds`, and one instance of `Wind Shear`. Remember, order doesn't matter in dictionaries.

First, you need to iterate through the file until you find an entry with a matching location name. And then?

A good tip is to try and perform the algorithm by hand first, and notice how you're iterating through each layer. It's ok if your code looks a bit ugly and heavily nested! Lots of code that works with real data is heavily nested too.

```python
# adverse event count

def get_adverse_event_count(airport_list, airport_name):
    adverse_events_dict = {}

    # Note! It's ok if your code looks a bit ugly!
    # Lots of code that works with real data is heavily nested!

    ...

    return adverse_events_dict

get_adverse_event_count(airports, "KBOS")
```

```python
grader.check("q_adverse_event_count")
```

## **3. Other Python Potpourri**

Alright, back to some more Python bits and bobs.

### **3.1 Tuples**

A **tuple** consists of a number of values separated by commas. They are much like a list, but are **immutable**. If you want to edit a value in a tuple, you would have to create a new tuple from scratch.

```python
data_tuple = (1, 2, 3)
list_item = ["A", "B", "C"]

print(data_tuple[2])
print(list_item[2])
```

```python
# You'll notice that if you try to edit the items of a tuple, you'll get an error
# Uncomment this line and try to run it

#data_tuple[1] = 3
```

Why would we use tuples over a list then? 

One of the most frequent places you'll see tuples is that it's the default return type of *functions that return multiple values*.

```python
def sum_and_product(a, b):
    s = a+b
    p = a*b
    return s, p

# Calling it returns a tuple
result = sum_and_product(4,2)

print(result)

# Getting the items from the tuple into their own variables
sum_item = result[0]
product_item = result[1]

print("Sum item:", sum_item)
print("Product item:", product_item)
```

In that last part: where we got the result of the function and split it into variables - that was called `unpacking`.

Any tuple can be unpacked into variables like this.

```python
a, b, c = (1, 2, 3)
print(a)
print(b)
print(c)
```

And when calling functions that return tuples and you want to put them into multiple variables, it'll usually just look like this.

```python
sum_item, product_item = sum_and_product(4, 8)
print(sum_item)
print(product_item)
```

#### **Exercise 3.1: Quadratic Formula.**

Write a function `quadratic_solver(a, b, c)` that returns both roots of $ax^2 + bx + c = 0$ using the quadratic formula.

As a reminder:

$x_1=\frac{-b + \sqrt{b^2-4ac}}{2a}$, $x_2=\frac{-b - \sqrt{b^2-4ac}}{2a}$

Assume that any coefficients given produce a real solution (no need to worry about negative square roots).

Python's [math](https://docs.python.org/3/library/math.html) library has any of the standard functions you might need.


Order matters! Return them as a tuple in this exact order: `x_1`, then `x_2`.

```python
import math

def quadratic_solver(a, b, c):
    ...

    return x_1, x_2
```

```python
grader.check("q_quadratic_solver")
```

### **3.2 Advanced Functions**

While we're on the topic of functions, let's talk about **keyword arguments**.

You may have used function keyword arguments before. They allow you to define functions where some parameters are optional and have default values.

```python
def trip_cost(distance_km, fuel_price=1.60, efficiency_kmpl=12.0, tolls=0.0):
    return (distance_km / efficiency_kmpl) * fuel_price + tolls

print(trip_cost(300))                                      
print(trip_cost(300, tolls=15.50))                       
print(trip_cost(300, efficiency_kmpl=8.0, fuel_price=1.85)) # Order doesn't matter for keyword args!
```

Finally, if we get real meta with it, Python also allows you to pass entire *functions* as parameters. You'll see more of this next week!

```python
def combine_and_square(a, b, combination_function):
    return combination_function(a, b) ** 2

def adder(a, b):
    return a + b

def multiplier(a, b):
    return a * b

def bigger(a, b):
    return max(a, b)

print(combine_and_square(3, 4, adder))       # (3+4)^2 = 49
print(combine_and_square(3, 4, multiplier))  # (3*4)^2 = 144
print(combine_and_square(4, 7, bigger))      # (7)^2 = 49

```

#### **Exercise 3.2: Keyword Arguments.** 

Write a function `filter_list(input_list, evens)`, where `input_list` is a list of numbers and `evens` is an *optional* parameter.

If the function is called with only `filter_list(input_list)`, it should by default return a list with only the even numbers in the list.

If the optional parameter `evens` is set to `False`, then it will return a list with only the odd numbers in the list.

If the list is empty, it should return an empty list.

```python
...

filter_list([1,2,3,5,6], False)
```

```python
grader.check("q_filter_list")
```

### **3.3 Sets**

Finally, one more thing that's good to know about is Sets.

Sets are like dictionaries but only store keys (no values), and every key is unique. They're great if your problem involves something along the lines of "have I seen this before?"

```python
basket_set = set()

basket_set.add("apple")
basket_set.add("orange")
basket_set.add("pear")

print(basket_set)

# Just like dicts: keys are unique
basket_set.add("pear")
print(basket_set)
```

```python
# You can also easily create a set from a list
# But remember, just like dicts, no duplicates!
new_set = set(["The", "Cat", "In", "The", "Mat", "In", "The", "House"])
new_set
```

The nice things about sets are that they also come with common set operations built in.

```python
in_cs_31 = set(["John", "Jame", "Jeanie", "Xu"])
in_cs_35 = set(["Xu","John","Bob","Sarah"])

# Who's in both classes? (Set Intersection)
print(in_cs_31 & in_cs_35)

# Who's in CS_31 but not in CS_35?
print(in_cs_31 - in_cs_35)

# Who's in 31, 35, or both? (Set Union)
print(in_cs_31 | in_cs_35)
```

## **4. Classes and Objects**

[Classes](https://www.w3schools.com/python/python_classes.asp) are used to bundle data and functionality that logically belongs together. Think of them as a template or a blueprint. Classes have fields (or properties) that determine what data belongs in the class. You use classes to create *instances* of a classes - called objects. 

For example, a class representing a person may have fields such as name, age, height, etc. Instances of these classes would be actual people with that data filled in.

Oftentimes you will be *using* classes from existing libraries. So it's good to know how they're constructed.

I've commented the structure of the basic properties of a class below.

```python
class Person: # Name of the class

  # __init__ is called the *constructor*
  # This is always called when an object is created
  def __init__(self, name, age): 
    # Fields of class
    self.name = name 
    self.age = age 

  # A method belonging to the class
  def introduce(self): 
    print("Hi! My name is "+self.name)
    print("I am", self.age, "years old.")
```

```python
p1 = Person("Emil", 36)

# You access the fields of the class via "."
print(p1.name)
print(p1.age) 

# Same thing with calling functions from within a class
# Functions that are inside a class are sometimes called "methods"
p1.introduce()
```

```python
p2 = Person("Niko", 42)
p2.introduce()
```

All classes have a built-in method called `__init__()`, which is always executed when the class is being initiated. This is called the *constructor*.

The `__init__()` method is used to assign values to object properties, or to perform operations that are necessary when the object is being created.

The `self` parameter is a reference to the current instance of the class, and must be the first parameter of any method in the class.

Let's create a new `MeetablePerson` class that allows one `MeetablePerson` to meet another `MeetablePerson`.

```python
class MeetablePerson:
    def __init__(self, name, age):
        self.name = name
        self.age = age

    def introduce(self):
        print("Hi! My name is "+self.name)
        print("I am", self.age, "years old.")

    def say_hi_to(self, other_person):
        print(self.name, "says hi to", other_person.name)
```

```python
p1 = MeetablePerson("Emil", 22)
p2 = MeetablePerson("Niko", 33)
p1.introduce() 
p1.say_hi_to(p2)
```

We can modify the fields in a class just like how we modify normal variables.

Let's create a class called `AgingPerson`, such that when the function `celebrate_birthday()` is called, the `AgingPerson`'s `age` field increases by one.

```python
class AgingPerson:
  def __init__(self, name, age):
    self.name = name
    self.age = age

  def celebrate_birthday(self):
    self.age += 1
    print(f"Happy birthday {self.name}! You are now {self.age}")

p1 = AgingPerson("Linus", 25)
print(p1.age)
p1.celebrate_birthday()
p1.celebrate_birthday() 

p2 = AgingPerson("Joe", 28)
p2.celebrate_birthday()
```

#### **Exercise 4.1: The `Robot` Class.** 

Define a class called `Robot` that has the fields `x_position` (int), `y_position` (int), and `speed` (int).

The `Robot` class should have the following functions. If no x or y parameters are given, the robot should default to (0,0) with a speed of 1

- `get_position()`: Returns a tuple containing the `x_position` and the `y_position`.
- `step_direction(direction)`: Takes in a string `direction` that is either `"up", "down", "left", or "right"`.
    - When step_direction(direction) is called, it should modify either the `x_position` or the `y_position` of the robot by the robots `speed` to represent the robot moving in that direction.
    - For example, `step_direction("right")` with a speed of `1` will increase `x_position` by 1.
    - `step_direction("up")` with a speed of `2` will increase `y_position` by 2.
- `check_collide(other_robot)`: Checks to see if the `other_robot` has the same location as this robot.

Don't forget to include the `self` param!

```python
class Robot:
    ...


# Some code to test your implementation

r1 = Robot(0,0,2)
r2 = Robot(4, 2)

print(r1.get_position())
print(r1.check_collide(r2)) # Should return "False"
r1.step_direction("up")
print(r1.get_position())
r1.step_direction("right")
r1.step_direction("right")
print(r1.get_position())

print(r1.check_collide(r2)) # Should return "True"
```

```python
grader.check("q_robot")
```

## **5. Data Structures and Complexity**

Good news - you're already familiar with using data structures! You've written code that has used `lists` and `dictionaries`.

However, there are many [*many*](https://en.wikipedia.org/wiki/List_of_data_structures) more types of data structures.

Usually, the first step to creating efficient algorithms that work on data is to be able to efficiently *access* data - and depending on the type of problem, it lends itself naturally to a different type of data structure.


### **5.1 Lists vs. Dicts: A Complexity Perspective**

As an illustrative example - let's consider `lists` vs. `dictionaries`.

Lists are one of the most basic data structures, that allow you to package multiple items together in a specific order. Compare this to dictionaries, which you use to store **key-value** pairs. Given a key, you can easily look up its corresponding value.
What is really nice is that given a key, looking up its value is a **constant time operation**.

To demonstrate how big of a deal this is, let's pretend that dictionaries didn't exist in Python, but you still wanted something like it. We can implement a naive version of a dictionary that searches through an entire list until it finds the item it's looking for. This is a *linear scan*.


```python
class NaiveDictionary:
    def __init__(self):
        self.data = []

    def add_kv_pair(self, key, value):
        self.data.append((key, value))
        return 0

    def lookup(self, key):
        i = 0
        while i < len(self.data):
            if self.data[i][0] == key:
                return self.data[i][1] # returns the value
            i+=1
```

```python
naive_dict = NaiveDictionary()
DICT_SIZE = 10**5 # We can do exponents by using **. This is 10^5

import time

for i in range(0, DICT_SIZE):
    # The first element in the tuple is our key, the second is the value
    naive_dict.add_kv_pair(i, i**2)

# Call the lookup method to scan the whole list for a key.
start = time.time()
naive_dict.lookup(DICT_SIZE//2)
end = time.time()

print("Naive list dictionary lookup time: ", end - start, "seconds")
```

Now, let's use just a normal dictionary to solve the problem.

```python
real_dict = {x: x**2 for x in range(DICT_SIZE)}

start = time.time()
real_dict[DICT_SIZE//2]
end = time.time()

print("Real dictionary: ", end - start, "seconds")
```

#### **Exercise 5.1 (Written): Lookup Speed Comparison.** **(15 points)**

Run your code for different values of `DICT_SIZE`: $10^3$, $10^4$, $10^5$, $10^6$ and $10^7$. Compare the lookup speeds for both the naive dictionary lookup, and the real dictionary lookup. Report your findings.

While this is a somewhat contrived example - this illustrates why careful data structure selection can vastly influence performance.

<!-- BEGIN QUESTION -->



```python
# Scratch space: try DICT_SIZE = 10**3, 10**4, 10**5, 10**6, 10**7 and time both
# naive_dict.lookup(...) (the naive list-based lookup) and real_dict[...] lookups.

```

**Your findings:**

_Type your answer here, replacing this text. Double click on this text to make edits._

<!-- END QUESTION -->

### **5.2 Stacks and Queues**

In many types of problems, we often want to iterate though a collection of objects in a set order.

Python wraps both of these up into a class called `deque` (pronounced “deck” and is short for “double-ended queue”). Since deques allow fast access to both elements at the front as well as the back, they can function as both a stack and a queue.

<center><img src="https://media.geeksforgeeks.org/wp-content/uploads/20250910175031267287/Deque.webp"></center>


#### **Stacks**

Stacks are **First In Last Out** data structures. Much like a stack of paper or plates.

```python
from collections import deque 
```

```python
# Making a stack using a deque

stack = deque()

stack.append(1)
stack.append(2)
stack.append(3)

print(stack)

# .pop() returns the element popped off the stack
first_popped = stack.pop()
print(first_popped)
print(stack)

print(stack.pop())
print(stack)

print(stack.pop())
print(stack)
```

```python
# Notice that if we pop from an empty deque, we get an error
empty_deque = deque()

# Uncomment the line below and see what happens.
# empty_deque.pop()
```

#### **Pause and Ponder**

Before running the following code cell, what do you think it will print?

```python
stack = deque()

stack.append(1)
stack.append("John")
stack.append(4.5233)
stack.append(True)

stack.pop()
stack.pop()
element = stack.pop()
print(stack)
print(element)
```

#### **Queues**

Queues are **First In First Out** data structures. Much like, well, a queue!

```python
queue = deque()

queue.append("A")
queue.append("B")
queue.append("C")

# .popleft() returns the element popped off the queue
first_popped = queue.popleft()
print(first_popped)
print(queue)

print(queue.popleft())
print(queue)

print(queue.popleft())
print(queue)
```

#### **Pause and Ponder**

Before running the following code cell, what do you think it will print?

```python
queue = deque()

queue.append("A")
queue.append("B")
queue.append("C")
queue.append("D")

while len(queue) != 0:
    print("Next element: "+queue.popleft())

queue.append("E")

print(queue)
```

#### **Exercise 5.2: `execute_move_queue`.**

Recall the `Robot` class you created earlier.

Create a function called `execute_move_queue(robot, queue)` which takes in a `Robot` and a `deque` of strings, each string being a move ("right", "left", "down", or "up").

When `execute_move_queue` is called, the robot should perform all of the moves in the queue, in order.

For example:

If our `Robot` started at `0,0` with a speed of `1`, and had a move queue of `["right","right","left","right","up"]`, the final position of the `Robot` should be `(2,1)`.

```python
directions = ["right","right","left","right","up"]

dir_queue = deque(directions)

def execute_move_queue(robot, queue):
    ...

start_robot = Robot()
start_robot.get_position()

execute_move_queue(start_robot, dir_queue)
start_robot.get_position()
```

```python
grader.check("q_execute_move_queue")
```

## **6. Putting It All Together**

Let's put everything we've learned together.

You've just started your internship for *CascadiaAir*, a regional airline, and they're trying to improve the efficiency of their dispatch system by looking at some models and simulations of the departure order of airplanes from the airfield. They've hired you to continue working on the project, and you've been handed the following code from the prior intern who is no longer there.

Below is the class structure from the data that is loaded from `airports.json` (which you've seen earlier), but also `flights.json`.

You should view `airports.json`, `flights.json`, and the classes below to see how they are structured.

Further below that is a function called `modify_priority`, which takes in an airplane and then modifies its priority in the dispatch order following a set of rules. Some rules are already implemented, and you will be asked to implement two more.



#### **Some Recommendations and Tips Before Getting Started**
 - **DON'T JUMP RIGHT INTO WRITING CODE**
 - You should read the entirety of the class structure, as well as the Queue before making any changes.
 - Diagram out (on paper, on a whiteboard, etc) the different classes, their fields, and how they relate to each other.
 - Consider making simple instances of each class to see how they work.
 - When making changes, write a little, then test a little. Don't try and do everything all at once.
 - When reading through code, write your own comments tracing your own understanding of how data flows.
 - Take breaks - understanding large amounts of existing code can be exhausting!

#### **Exercise 6.1: Priority Assignment.** 

As you read through the `ForecastEntry`, `Forecast`, `Airport`, and `Flight` classes below, you'll notice two fields aren't filled out yet: `has_heavy_weather` on `Forecast`, and `is_high_capacity` on `Airport`. These are calculated fields that will help you make decisions on assigning priority. You'll then use both of them inside `modify_priority()` further down.

Here are the rules for the dispatch queue. Rules 0 and 1 have already been implemented. Your job is to implement rules 2 and 3.


- **Rule 0 (Already Implemented)**: 
    - Other rules withstanding, planes should be assigned a priority based on their designation:
    - Medivac > VIP > Commercial > Light Aircraft (e.g. All medivacs depart before any VIPs)
    - The prior intern assigned priorities to numbers as follows:
        - 0: Medivac
        - 1: VIP
        - 2: Commercial
        - 3: Light Aircraft
        - 4: Do Not Fly
- **Rule 1 (Already Implemented)**:  
    - If at the plane's destination, the projected cumulative rainfall is greater than the plane's max_rain threshold, then they are not clear to fly.
- **Rule 2 (You Implement)**: 
    - If the forecast has projected "Heavy Weather" for a plane's destination, then **Light Aircraft** are not cleared to fly. All other aircraft are clear to fly, regardless of heavy weather.
    - A forecast is considered to have heavy weather **if more than half** of the forecast entries have at least one adverse weather events present. Fill this in as `has_heavy_weather` on `Forecast`.
- **Rule 3 (You Implement)**: 
    - If a destination is at high capacity, incoming aircraft should not be cleared to fly, UNLESS it is a Medivac.
    - A destination is considered high capacity if its current occupancy is >90% of its max capacity. Fill this in as `is_high_capacity` on `Airport`.

Reference both fields inside `modify_priority()` to enforce Rules 2 and 3, the same way the rain rule (Rule 1) already does using `cumulative_rain`.

```python
class ForecastEntry:
    def __init__(self, temperature, rain_level, adverse_events):
        self.temperature = temperature          # float, degrees C
        self.rain_level = rain_level            # float, rainfall for this entry
        self.adverse_events = adverse_events    # list[str], empty list if none

    def has_adverse_event(self):
        """True if this entry lists at least one adverse weather event."""
        return len(self.adverse_events) > 0

    def __repr__(self):
        return (f"ForecastEntry(temp={self.temperature}, "
                f"rain={self.rain_level}, events={self.adverse_events})")

class Forecast:
    """
    Wraps a list of ForecastEntry objects.

    Note the two attributes calculated in __init__. We work them out ONCE,
    when the data is loaded, instead of recalculating them every time a plane
    asks about the weather. Keep this in mind for the complexity question.
    """

    def __init__(self, entries):
        self.entries = entries                              # list[ForecastEntry]

        # Total projected rainfall = the sum of every entry's rain_level.
        self.cumulative_rain = sum(e.rain_level for e in entries)

        # IMPLEMENT RULE 2 HERE 
        
        ...

        self.has_heavy_weather = ...

    def __len__(self):
        return len(self.entries)

    def __repr__(self):
        return (f"Forecast({len(self.entries)} entries, "
                f"cumulative_rain={self.cumulative_rain:.1f}, "
                f"heavy_weather={self.has_heavy_weather})")
```

```python
class Airport:

    def __init__(self, location, city, max_capacity, current_occupancy, forecast):
        self.location = location                # str, e.g. "KBOS"
        self.city = city                        # str, e.g. "Boston"
        self.max_capacity = max_capacity        # int, total parking spots
        self.current_occupancy = current_occupancy   # int, spots currently USED
        self.forecast = forecast                # Forecast object

        # IMPLEMENT RULE 3 HERE
        self.is_high_capacity = ...


    def __repr__(self):
        return (f"Airport({self.location}, {self.city}, "
                f"{self.current_occupancy}/{self.max_capacity}, "
                f"{self.forecast!r})")

class Flight:
    # Lower number = higher priority. Medivac departs before VIP, and so on.
    PRIORITY_ORDER = {
        "Medivac": 0,
        "VIP": 1,
        "Commercial": 2,
        "Light Aircraft": 3,
    }

    def __init__(self, flight_id, designation, destination, max_rain_thresh):
        self.flight_id = flight_id              # str, e.g. "NMED001"
        self.designation = designation          # str, one of PRIORITY_ORDER
        self.destination = destination          # str, an Airport's location code
        self.max_rain_thresh = max_rain_thresh  # float, max rain it can land in
        self.priority = Flight.PRIORITY_ORDER[self.designation]

    def get_priority(self):
        return self.priority

    def __repr__(self):
        return (f"Flight({self.flight_id}, {self.designation}, "
                f"-> {self.destination})")



```

```python
def load_airports(path="airports.json"):
    """Read airports.json and return a list of Airport objects."""
    with open(path) as f:
        raw_airports = json.load(f)

    airports = []
    for raw in raw_airports:
        # Build a ForecastEntry for each dictionary in the "forecast" list...
        entries = [
            ForecastEntry(
                temperature=e["temperature"],
                rain_level=e["rain_level"],
                adverse_events=e["adverse_events"],
            )
            for e in raw["forecast"]
        ]
        # ...then wrap them all in one Forecast object.
        forecast = Forecast(entries)

        airports.append(Airport(
            location=raw["location"],
            city=raw["city"],
            max_capacity=raw["max_capacity"],
            current_occupancy=raw["current_occupancy"],
            forecast=forecast,
        ))
    return airports


def load_flights(path="flights.json"):
    """Read flights.json and return a list of Flight objects."""
    with open(path) as f:
        raw_flights = json.load(f)

    return [
        Flight(
            flight_id=raw["flight_id"],
            designation=raw["designation"],
            destination=raw["destination"],
            max_rain_thresh=raw["max_rain_thresh"],
        )
        for raw in raw_flights
    ]

def build_airport_lookup(airports):
    """
    Build a dictionary that maps a location code -> Airport object.
    We do this for quick O(1) access to different airports.
    """
    return {airport.location: airport for airport in airports}
```

```python
airport_objects = load_airports("airports.json")
flights = load_flights("flights.json")
airports_by_location = build_airport_lookup(airport_objects)

print(f"Loaded {len(airport_objects)} airports and {len(flights)} flights.")
```

```python
print("Example airport:", airport_objects[0])
print("Example flight: ", flights[0])
print()
print("Looking up that flight's destination:")
destination = airports_by_location[flights[0].destination]
print("  ", destination)
```

```python
def modify_priority(airplane):
    # Get our destination airport for our current flight
    location = airports_by_location[airplane.destination]

    # Reset to the base priority for this flight so it can be re-run safely.
    airplane.priority = Flight.PRIORITY_ORDER[airplane.designation]

    # RULE ONE (ALREADY IMPLEMENTED)
    if location.forecast.cumulative_rain > airplane.max_rain_thresh:
        airplane.priority = 4



    # IMPLEMENT RULE 2 HERE

    ...

    # IMPLEMENT RULE 3 HERE

    ...

    return airplane
```

The prior intern made a class called `AirplaneQueue`, which provides an interface for adding and removing airplanes from the queue following the priority rules - with his comments.

```python

class AirplaneQueue:
    def __init__(self):
        # One list holding every waiting plane, kept in departure order.
        self.queue = []

    # TODO: It works, but it's really slow for a lot of planes
    # I think there's gotta be a faster way to do this...?

    def add_to_queue(self, airplane):
        airplane_update = modify_priority(airplane)
        new_priority = airplane_update.get_priority()
 
        # Walk the list from the front until we reach the first plane that
        # this one should depart BEFORE, then insert into the correct
        # queue location.
        insert_at = len(self.queue)          # default: the very back

        for i in range(len(self.queue)):
            if self.queue[i].get_priority() > new_priority:
                insert_at = i
                break
 
        self.queue.insert(insert_at, airplane_update)

    # Remove the plane from the front
    def remove_from_queue(self):
        popped_item = None
 
        if len(self.queue) != 0:
            popped_item = self.queue.pop(0)
 
        return popped_item
 
    def __len__(self):
        return len(self.queue)

```

```python
air_queue = AirplaneQueue()


air_queue.add_to_queue(flights[30])
air_queue.add_to_queue(flights[10])

print(air_queue.remove_from_queue())
```

```python
grader.check("q_priority_rules")
```

#### **Exercise 6.2: `FastAirplaneQueue`.**

Now that you know a bit more about data structures, you get the feeling that the code for managing the priority queue is inefficient. You ask your friend in Computer Science over lunch who took Data Structures and Algorithms last semester for some help. She mentions for problems like this where there are a fixed number of priority levels, one of the most efficient approaches is to have an underlying data structure that is a dictionary of queues - one for each priority level.

You hand her your laptop and she begins to type up an implementation, but before she can fully explain how it works and why it's more efficient, lunch is over and she has to go!

**Implement both the `add_to_queue` and `remove_from_queue` functions for this new data structure in `FastAirplaneQueue`.** Remember, the end result of the order of things popped should be exactly the same as `AirplaneQueue`.

```python
class FastAirplaneQueue:

    def __init__(self):
        self.queues = {
            0: deque(), # Medivac - HIGHEST PRIORITY
            1: deque(), # vip
            2: deque(), # commercial
            3: deque(), # light aircraft
            4: deque(), # do not fly - LOWEST PRIORITY
        }

    def add_to_queue(self, airplane):
        ...

    def remove_from_queue(self):
        popped_item = None

        ...

        return popped_item

```

```python
grader.check("q_fast_airplane_queue")
```

#### **Exercise 6.3 (Written): Complexity.** **(20 points)**

Using what we talked about complexity analysis, explain why the multiple queues approach (`FastAirplaneQueue`) would be more efficient than the prior approach (`AirplaneQueue`). Determine the complexity of `AirplaneQueue` as well as `FastAirplaneQueue` and reference them in your answer. You will need to think about the complexity of common `deque` operations, as well as the complexity of iterating through a list.

**Optional Challenge (no extra credit, for your own curiosity):** This approach is more efficient for a *fixed* number of priorities (e.g. in our case where it was 0-4).  Why *wouldn't* this be more efficient, complexity wise, if we had an unbounded number of priorities (e.g. can assign any integer priority to a plane)?

<!-- BEGIN QUESTION -->

_Type your answer here, replacing this text. Double click on this text to make edits._

<!-- END QUESTION -->

#### **Exercise 6.4 (Written): Wait, Is This AI?** **(15 points)**

Think back to the first week when we discussed the bigger question of what constitutes AI. In your own words - do you consider this airplane queueing system you wrote to be an instance of AI? If so, explain why, and give examples of changes that you feel would make it not AI. If not, explain why not - and give examples of changes that *would* make it AI. You must explicitly reference one of the concepts or definitions we talked about in-class in your discussion. Your response should be no more than a paragraph or two.

<!-- BEGIN QUESTION -->

_Type your answer here, replacing this text. Double click on this text to make edits._

<!-- END QUESTION -->

## **Collaboration Statement** **(10 points)**

You are required to complete the following collaboration statement.

<!-- BEGIN QUESTION -->

- If you completed the lab in a group - list all your group members here, including yourself. **If you worked individually, you must state *"I worked by myself."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- You must list the names of any others you got help from when completing this assignment (including TAs), and in what way you worked with them (discussed ideas, debugged math). **If you did not collaborate at all, you must state *"I did not collaborate with anyone else."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- If you found any resources online helpful, you must link them and describe how you found them helpful (or not helpful). **If you did not use any resources online, you must state *"I did not use any resources online."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- If you used any AI assistive technologies, you must describe how you used them and how you found them helpful (or not helpful). **If you did not use any AI assistive technologies, you must state *"I did not use any AI assistive technologies."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- If you have any questions or comments about the material, or if there was a topic or question you were confused about - please make a note of it here! **If you do not have any questions or comments, you must state *"I do not have any questions or comments about the assignment."***
    - _Type your answer here, replacing this text. Double click on this text to make edits._
- Estimate how long this assignment took you in hours. **You must give an estimate**.
    - This assignment took me about _n_ hours to complete.

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Run every cell from top to bottom and SAVE before exporting. Then submit
the generated .zip to the HW2 autograder assignment on Gradescope.


```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q_adverse_event_count (points: None)
```
>>> assert get_adverse_event_count(airports, 'KBOS') == {'Snow': 1, 'High Winds': 2, 'Wind Shear': 1}
```
```
>>> assert get_adverse_event_count(airports, 'KPHX') == {'Wind Shear': 2, 'Hail': 1}
```
```
>>> assert get_adverse_event_count(airports, 'NOT_A_REAL_AIRPORT') == {}
```
failure_message: An airport with no matching location should give back an empty dict.


Tests for q_dict_filter (points: None)
```
>>> assert new_dict == {'Bob': 1, 'Alice': 2}
```
```
>>> assert d == {'Bob': 1, 'Alice': 2, 'Zuzu': 4, 'Blarg': 9}
```
failure_message: d should be left unchanged by your comprehension.


Tests for q_execute_move_queue (points: None)
```
>>> r = Robot()
>>> q = deque(['right', 'right', 'left', 'right', 'up'])
>>> execute_move_queue(r, q)
>>> assert r.get_position() == (2, 1)
```
```
>>> r = Robot()
>>> q = deque(['up', 'up', 'down'])
>>> execute_move_queue(r, q)
>>> assert len(q) == 0 and r.get_position() == (0, 1)
```
failure_message: Make sure your queue ends up empty (moves are consumed, not just read).


Tests for q_fast_airplane_queue (points: None)
```
>>> q = FastAirplaneQueue()
>>> cmr = next((f for f in flights if f.designation == 'Commercial' and f.destination == 'KBOS'))
>>> q.add_to_queue(cmr)
>>> assert q.remove_from_queue() is cmr
```
```
>>> q = FastAirplaneQueue()
>>> by_tier = {}
>>> for desig in ['Light Aircraft', 'Commercial', 'VIP', 'Medivac']:
...     f = next((fl for fl in flights if fl.designation == desig and fl.destination == 'KBOS'))
...     q.add_to_queue(f)
>>> order = [q.remove_from_queue().designation for _ in range(4)]
>>> assert order == ['Medivac', 'VIP', 'Commercial', 'Light Aircraft']
```
```
>>> q = FastAirplaneQueue()
>>> same_tier = [Flight(f'TEST{i}', 'Commercial', 'KPHX', 999.0) for i in range(3)]
>>> for f in same_tier:
...     q.add_to_queue(f)
>>> popped = [q.remove_from_queue() for _ in range(3)]
>>> assert popped == same_tier
```
failure_message: Check that you're using popleft(), not pop().
```
>>> q = FastAirplaneQueue()
>>> blocked = next((f for f in flights if f.designation == 'Light Aircraft' and f.destination == 'KJFK'))
>>> cmr = next((f for f in flights if f.designation == 'Commercial' and f.destination == 'KBOS'))
>>> q.add_to_queue(blocked)
>>> q.add_to_queue(cmr)
>>> assert q.remove_from_queue() is cmr
>>> assert q.remove_from_queue() is blocked
```


Tests for q_filter_list (points: None)
```
>>> assert filter_list([1, 2, 3, 5, 6], False) == [1, 3, 5]
```
```
>>> assert filter_list([1, 2, 3, 4, 5, 6]) == [2, 4, 6]
```
failure_message: Check the default value of evens.
```
>>> lst = [4, 1, 6, 3, 2, 5]
>>> result = filter_list(lst)
>>> assert lst == [4, 1, 6, 3, 2, 5]
```


Tests for q_odds_squared (points: None)
```
>>> assert odds_squared == [1, 9, 25, 49, 81]
```
```
>>> assert nums == list(range(1, 11))
```
failure_message: nums should be left unchanged by your comprehension.


Tests for q_priority_rules (points: None)
```
>>> assert airports_by_location['KORD'].is_high_capacity is True and airports_by_location['KBOS'].is_high_capacity is False
```
failure_message: is_high_capacity is not implemented correctly.
```
>>> assert airports_by_location['KJFK'].forecast.has_heavy_weather is True and airports_by_location['KBOS'].forecast.has_heavy_weather is False
```
failure_message: has_heavy_weather is not implemented correctly.
```
>>> med_rain_blocked = next((f for f in flights if f.designation == 'Medivac' and f.destination == 'KDEN' and (airports_by_location['KDEN'].forecast.cumulative_rain > f.max_rain_thresh)))
>>> assert modify_priority(med_rain_blocked).priority == 4
>>> med_cap_exempt = next((f for f in flights if f.designation == 'Medivac' and f.destination == 'KORD'))
>>> assert modify_priority(med_cap_exempt).priority == 0
```
failure_message: A Medivac should still be blocked by the rain rule - the capacity exemption only applies to Rule 3.
```
>>> light_heavy_wx = next((f for f in flights if f.designation == 'Light Aircraft' and f.destination == 'KJFK'))
>>> assert modify_priority(light_heavy_wx).priority == 4
```
failure_message: A light aircraft was not blocked by heavy weather.


Tests for q_quadratic_solver (points: None)
```
>>> x1, x2 = quadratic_solver(1, -8, 5)
>>> assert math.isclose(x1, 7.3166247903554, rel_tol=1e-09) and math.isclose(x2, 0.6833752096446002, rel_tol=1e-09)
```
```
>>> x1, x2 = quadratic_solver(2, -8, 6)
>>> assert math.isclose(x1, 3.0, rel_tol=1e-09) and math.isclose(x2, 1.0, rel_tol=1e-09)
```
failure_message: Double check your parentheses around the denominator 2*a.


Tests for q_robot (points: None)
```
>>> r1, r2 = (Robot(0, 0, 2), Robot(4, 2))
>>> assert r1.get_position() == (0, 0) and r1.check_collide(r2) is False
```
```
>>> r = Robot(0, 0, 2)
>>> r.step_direction('up')
>>> r.step_direction('right')
>>> r.step_direction('right')
>>> assert r.get_position() == (4, 2)
```
```
>>> assert Robot().get_position() == (0, 0) and Robot().speed == 1
```
failure_message: Check your default arguments for x_position, y_position, and speed.
```
>>> a, b = (Robot(2, 2, 1), Robot(2, 2, 5))
>>> assert a.check_collide(b) is True
>>> assert Robot(3, 0).check_collide(Robot(3, 5)) is False
>>> assert Robot(2, 5).check_collide(Robot(8, 5)) is False
```
```
>>> r = Robot(5, 5, 3)
>>> r.step_direction('left')
>>> r.step_direction('down')
>>> assert r.get_position() == (2, 2)
```
failure_message: Check your left/down branches in step_direction.
