```python
# Initialize Otter
import otter
grader = otter.Notebook("otter_demo.ipynb")
```

# HW0: Demo Lab
#### CS-31: Foundations for Artificial Intelligence

*You are allowed to work in groups for the lab. You must write anything that will be turned in -- all code and all written solutions -- on your own. You must list all people and resources that you interacted with in the Collaboration Statement below. See [course policies](https://www.cs.tufts.edu/comp/31/policy.html) for more information.*

--

In this assignment...


We understand that you may not have programmed in a while, or that you've programmed before but not in Python. It will probably take time to remember some of the basics and get back in the groove, and that's ok! **However if after some time you find yourself severely struggling with the problems below, please see the instructor.**

### **Table of Contents**

0) Basic Notebook Navigation
1) Notebooks and Python

## **0. Basic Navigation and Autograder** 

Before we even start, doing these two things will help make navigating Jupyter Notebooks much easier and reduce clutter:

1) Near the top of this window, you'll see a button that says **Outline** (right next to **Jupyter Variables**). Click it.
    - You'll see a window pop up that shows a navigable outline for the notebook that you can keep handy.
2) Click on the text **0. Basic Navigation** right above this cell. You'll see a **⌄** icon right to the left. Click it, and it will collapse everything in this section. Right under it you'll see something like `1 cell hidden`. If you click it again you'll see a **>** symbol: click that again to expand the section.
    - Collapse all of the top level headers for sections 1, 2, and 3: and you can reveal them as you continue to keep the amount of text on the screen low.
    - You can do the same *within sections* (e.g. try collapsing and uncollapsing **0.1 Navigation Test** below) as you're working through them.


### **0.1 Navigation Test**

I'm hidden!

## **1. Notebooks and Python**

### **1.1 Cells**

Below is a *code cell*. These are where you do most of your computation.

If you click on the code cell, it will be highlighted on the left. You'll see a play button show up at the top left just outside of the cell, and two brackets (`[]`) show up below that. Sometimes you'll see this space to the left of the cell referred to as the *gutter*.

**Click the play icon to execute the cell**.

If you just opened up VSCode, you may see a window pop up near the top to select your Python environment. Select the CS31 environment. This is sometimes called the *kernel*.

When a cell is run, it automatically runs and then prints the last line in the cell. You'll also see a `[1]` show up in the bottom left. This is the *execution order* of the cell. You'll notice if you execute this cell multiple times in a row this will tick up.




```python
# Run me!

"Hello world!"
```

Output:
```
'Hello world!'
```

You can also run cells by highlighting it and pressing **Cmd/Ctrl+Enter** (my preferred method).

Running a cell evaluates all of the code inside of it.

```python
# Run me by pressing Cmd/Ctrl+Enter

9+10
```

Output:
```
19
```

If you want to run cells in a sequence, you can use **Shift+Enter** to run the cell and move focus to the next cell.

Try running the following cells in sequence using **Shift+Enter**.


```python
"Run me!"
```

```python
for i in range(1,10):
    print("Looping step", i)
```

Output:
```
Looping step 1
Looping step 2
Looping step 3
Looping step 4
Looping step 5
Looping step 6
Looping step 7
Looping step 8
Looping step 9
```

```python
for name in ["James", "Harry", "Chris"]:
    print(f"Hello there {name}!")
```

### **1.2. Variables, Functions and Persistence**

As you may remember, variables allow us to store and manipulate data.

```python
my_number = 5

print(my_number)
```

Output:
```
5
```

Now, what's both cool and tricky about notebooks is that variables *persist* between cells. 

```python
your_number = my_number + 5

print(your_number)
print(my_number)
```

Output:
```
10
5
```

In the top bar, click on the icon that says `Jupyter Variables` (next to clear all outputs).

You should see a tab open up at the bottom that showcases all of the variables in memory.

What variables do you see? 

These are all the variables that are currently loaded into memory that you can reference.

### **1.3 Functions**

Just like in normal `.py` files, you can define functions and then use them throughout your notebook.

```python
def get_greater_than(input_list, number):
    new_list = []

    for item in input_list:
        if item > number:
            new_list.append(item)

    return new_list
```

```python
item = [1,10,2,9,3,8,4,7,5,6]

get_greater_than(item, 5)
```

Output:
```
[10, 9, 8, 7, 6]
```

#### **Exercise 1.1: Adding Numbers**

**It's time to do your first Exercise!**

Write a function `add_numbers(a,b)` which adds up `a` and `b` and returns their sum. Then in the cell below that, create a variable called `add_20_30`, which should contain the sum of 20 and 30.

```python
def add_numbers(a, b):
    # BEGIN SOLUTION
    return a+b # SOLUTION
    # END SOLUTION
```

```python
# BEGIN SOLUTION
add_20_30 = add_numbers(20,30) # SOLUTION
# END SOLUTION
```

In this class we'll be using the `otter` autograder (that's what the cell at the very top of the notebook was for). For each exercise, it will run test cases to make sure you implemented your code correctly.

Make sure you run the cell at the top of the notebook, and then run the following cell. If your code is correct, you'll see a message that all tests were passed. If some of your code is wrong, you'll get an error message.

```python
grader.check("q_autograde")
```

Output:
```
q_autograde results: All test cases passed!
```

#### **Exercise 1.2 (Manual Code): Sample Manual and Written Question.** 

Write down some code to print out your name 10 times.

```python
# BEGIN SOLUTION
for i in range(10):
    print("Prof. Skripchuk")
# END SOLUTION
```

Output:
```
Prof. Skripchuk
Prof. Skripchuk
Prof. Skripchuk
Prof. Skripchuk
Prof. Skripchuk
Prof. Skripchuk
Prof. Skripchuk
Prof. Skripchuk
Prof. Skripchuk
Prof. Skripchuk
```

<!-- END QUESTION -->

#### **Exercise 1.3 (Manual): Sample Written Question.** 

Write down how you're feeling today.

<!-- BEGIN QUESTION -->

I am feeling good!

There is no right answer to how you're feeling today!

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Run every cell from top to bottom and SAVE before exporting. Then submit
the generated .zip to the HW1 autograder assignment on Gradescope.


```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q_autograde results: All test cases passed!
<IPython.core.display.HTML object>
```

Tests for q_autograde (points: None)
```
>>> assert add_numbers(10, 10) == 20
```
```
>>> assert add_20_30 == 50
```
